```python
# Initialize Otter
import otter

grader = otter.Notebook("fifa_21_raw_data.ipynb")
```

## Week 2 - FIFA '21 Player Ratings

This notebook is created for the week 2 assignment. Here, you have a very messy and raw dataset of EA Sports' installment of their hit FIFA series - FIFA21, which was scraped from sofifa.com.

Your task in this assignment is to clean up this dataset using the `Hints` provided in the ReadMe file

```python
import pandas as pd
import numpy as np
import os
%matplotlib inline
import otter
grader = otter.Notebook()
```

**Question 1:** Load the fifa dataset into a variable called `fifa_dataset`. Next, write a function called `check_data` to check if the `fifa_dataset` is empty or not by returning True or False.

```python
fifa_dataset = pd.read_csv('fifa_21_raw_data.csv', dtype={'Hits': str}, engine='python')

fifa_dataset.head()

def check_data(fifa_dataset):
   if fifa_dataset.empty:
      return True
   else:
      return False


```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 1:** Load the fifa dataset into a variable called `fifa_dataset`. Next, write a function called `check_data` to check if the `fifa_dataset` is empty or not by returning True or False.

**Question 2:** Remove all duplicate player `name` from the dataset. Remember, you need to leave the first occurence and remove other ones.

```python
def remove_duplicate(dataset):
    return dataset.drop_duplicates(subset=['Name'], keep='first')

```

```python
grader.check("q2")
```

**Question 1:** Load the fifa dataset into a variable called `fifa_dataset`. Next, write a function called `check_data` to check if the `fifa_dataset` is empty or not by returning True or False.

**Question 3:** Remove all leading white spaces or newline "\n" in the `club` column. For example, `"\n\n\n\nFC Barcelona"` should be transformed into `"FC Barcelona"`



```python
def remove_newline(dataset):
    dataset ['Club'] = dataset['Club'].str.strip()
    return dataset
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

**Question 4:** List the TOP 5 countries with the highest number of representation of players in the data. Create a dictionary called `top_5_country` using the country names as the key and the number of players as the value. For example,

`top_5_country` = {'Kenya': 15,
                  'Nigeria': 45,
                  'Ghana': 14,
                  'Canada': 20,
                  'USA': 88
                 }


```python
top_5_country = fifa_dataset['Nationality'].value_counts().head(5)

top_5_country = top_5_country.to_dict()

print(top_5_country)

fifa_dataset = remove_duplicate(remove_newline(fifa_dataset))

fifa_dataset.describe()

```

Output:
```
{'England': 1705, 'Germany': 1195, 'Spain': 1065, 'France': 1003, 'Argentina': 943}
                  ID           Age          ↓OVA           POT           BOV  \
count   17920.000000  17920.000000  17920.000000  17920.000000  17920.000000   
mean   225993.853292     25.244531     65.894141     71.234821     66.919308   
std     27165.550622      4.699357      6.981700      6.120768      6.758423   
min        41.000000     16.000000     47.000000     47.000000     48.000000   
25%    209607.500000     21.000000     61.000000     67.000000     63.000000   
50%    231863.500000     25.000000     66.000000     71.000000     67.000000   
75%    246456.250000     29.000000     70.000000     75.000000     71.000000   
max    259216.000000     53.000000     93.000000     95.000000     93.000000   

          Attacking      Crossing     Finishing  Heading Accuracy  \
count  17920.000000  17920.000000  17920.000000      17920.000000   
mean     249.822600     49.907087     45.999442         5
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**Question 5:** What insight can you derive from this data?

We've 17920 uique players from 18979
Our players are aged 25 on average with 16 years old being the younges and 53 being the oldest 
England seems to have the highest number of players, followed immediately by Germany and Spain

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
/Users/chukwuemekaokonkwo/Applications/IDS/ids-fifa-21-player-ratings/venv/lib/python3.14/site-packages/otter/check/notebook.py:494: UserWarning: Could not locate a PDF to include
  warnings.warn("Could not locate a PDF to include")
File ~/Applications/IDS/ids-fifa-21-player-ratings/venv/lib/python3.14/site-packages/nbconvert/exporters/pdf.py:166, in PDFExporter.run_latex(self, filename, raise_on_failure)
    163 def log_error(command, out):
    164     self.log.critical("%s failed: %s\n%s", command[0], command, out)
--> 166 return self.run_command(
    167     self.latex_command, filename, self.latex_count, log_error, raise_on_failure
    168 )

File ~/Applications/IDS/ids-fifa-21-player-ratings/venv/lib/python3.14/site-packages/nbconvert/exporters/pdf.py:120, in PDFExporter.run_command(self, command_list, filename, count, log_function, raise_on_failure)
    114     link = "https://nbconvert.readthedocs.io/en/latest/install.html#installing-tex"
    115     msg = (
    116         f"{c
```

Tests for q1 (points: 1)
```
>>> fifa_dataset = pd.read_csv('fifa_21_raw_data.csv', engine='python')
>>> 
>>> check_data(fifa_dataset)
False
```


Tests for q2 (points: 1)
```
>>> remove_duplicate(fifa_dataset)
>>> fifa_dataset['Name'].duplicated().any() == False
True
```


Tests for q3 (points: 1)
```
>>> len(fifa_dataset["Name"].unique()) > 1000
True
```
```
>>> type(fifa_dataset["Club"]) == pd.Series
True
```


Tests for q4 (points: 1)
```
>>> type(top_5_country) is dict
True
```
```
>>> len(top_5_country) == 5
True
```
