```python
# Initialize Otter
import otter
grader = otter.Notebook("fifa_21_raw_data.ipynb")
```

## Week 2 - FIFA '21 Player Ratings

This notebook is created for the week 2 assignment. Here, you have a very messy and raw dataset of EA Sports' installment of their hit FIFA series - FIFA21, which was scraped from sofifa.com.

Your task in this assignment is to clean up this dataset using the `Hints` provided in the ReadMe file

```python
import pandas as pd
import numpy as np
import os
%matplotlib inline
import otter
grader = otter.Notebook()
```

**Question 1:** Load the fifa dataset into a variable called `fifa_dataset`. Next, write a function called `check_data` to check if the `fifa_dataset` is empty or not by returning True or False.

```python
fifa_dataset = pd.read_csv('fifa_21_raw_data.csv', engine='python') # SOLUTION

def check_data(fifa_dataset):
    # BEGIN SOLUTION
    return fifa_dataset.empty
    # END SOLUTION
```

```python
grader.check("q1")
```

**Question 2:** Remove all duplicate player `name` from the dataset. Remember, you need to leave the first occurence and remove other ones.

```python
def remove_duplicate(dataset):
    # BEGIN SOLUTION
    duplicate = dataset['Name'].duplicated().any()
    if duplicate:
        dataset.drop_duplicates(subset=['Name'], inplace=True)
    # END SOLUTION
```

```python
grader.check("q2")
```

**Question 3:** Remove all leading white spaces or newline "\n" in the `club` column. For example, `"\n\n\n\nFC Barcelona"` should be transformed into `"FC Barcelona"`



```python
def remove_newline(dataset):
    # BEGIN SOLUTION
    dataset['Club'] = dataset['Club'].replace('\n\n\n\n', '', regex=True)
    return dataset
    # END SOLUTION
```

```python
grader.check("q3")
```

**Question 4:** List the TOP 5 countries with the highest number of representation of players in the data. Create a dictionary called `top_5_country` using the country names as the key and the number of players as the value. For example,

`top_5_country` = {'Kenya': 15,
                  'Nigeria': 45,
                  'Ghana': 14,
                  'Canada': 20,
                  'USA': 88
                 }


```python
top_5_country = dict(fifa_dataset['Nationality'].value_counts().head(5)) # SOLUTION
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

**Question 5:** What insight can you derive from this data?

_Type your answer here, replacing this text._

**SOLUTION:** England has the highest number of players in FIFA '21

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: 1)
```
>>> fifa_dataset = pd.read_csv('fifa_21_raw_data.csv', engine='python')
>>> 
>>> check_data(fifa_dataset)
False
```
```
>>> assert fifa_dataset['Name'][0] == 'L. Messi'
>>>
```


Tests for q2 (points: 1)
```
>>> remove_duplicate(fifa_dataset)
>>> fifa_dataset['Name'].duplicated().any() == False
True
```
```
>>> # for name in fifa_dataset['Name'].unique():
>>> # assert fifa_dataset['Name'].unique() == True
```


Tests for q3 (points: 1)
```
>>> len(fifa_dataset["Name"].unique()) > 1000
True
```
```
>>> type(fifa_dataset["Club"]) == pd.Series
True
```
```
>>> def is_white_spaces(fifa_dataset):
...     remove_newline(fifa_dataset)
...     for value in fifa_dataset['Club']:
...         if value.__contains__('\n'):
...             return False
...     return True
>>> 
>>> is_white_spaces(fifa_dataset) == True
True
```


Tests for q4 (points: 1)
```
>>> type(top_5_country) is dict
True
```
```
>>> len(top_5_country) == 5
True
```
```
>>> top_5_country['England'] >= 1500
True
```
