```python
# Initialize Otter
import otter

grader = otter.Notebook("fifa_21_raw_data.ipynb")
```

## Week 2 - FIFA '21 Player Ratings

This notebook is created for the week 2 assignment. Here, you have a very messy and raw dataset of EA Sports' installment of their hit FIFA series - FIFA21, which was scraped from sofifa.com.

Your task in this assignment is to clean up this dataset using the `Hints` provided in the ReadMe file

```python
import pandas as pd
import numpy as np
import os
%matplotlib inline
import otter
grader = otter.Notebook()
```

**Question 1:** Load the fifa dataset into a variable called `fifa_dataset`. Next, write a function called `check_data` to check if the `fifa_dataset` is empty or not by returning True or False.

```python
fifa_dataset = pd.read_csv('fifa_21_raw_data.csv', dtype={'Hits': str}, engine='python')

fifa_dataset.head()

def check_data(fifa_dataset):
   if fifa_dataset.empty:
      return True
   else:
      return False


```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 1:** Load the fifa dataset into a variable called `fifa_dataset`. Next, write a function called `check_data` to check if the `fifa_dataset` is empty or not by returning True or False.

**Question 2:** Remove all duplicate player `name` from the dataset. Remember, you need to leave the first occurence and remove other ones.

```python
def remove_duplicate(dataset):
    return dataset.drop_duplicates(subset=['Name'], keep='first')

```

```python
grader.check("q2")
```

**Question 1:** Load the fifa dataset into a variable called `fifa_dataset`. Next, write a function called `check_data` to check if the `fifa_dataset` is empty or not by returning True or False.

**Question 3:** Remove all leading white spaces or newline "\n" in the `club` column. For example, `"\n\n\n\nFC Barcelona"` should be transformed into `"FC Barcelona"`



```python
def remove_newline(dataset):
    string_cols = dataset.select_dtypes(include='str').columns
    dataset[string_cols] = dataset[string_cols].apply(lambda col: col.str.strip())
    return dataset
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

**Question 4:** List the TOP 5 countries with the highest number of representation of players in the data. Create a dictionary called `top_5_country` using the country names as the key and the number of players as the value. For example,

`top_5_country` = {'Kenya': 15,
                  'Nigeria': 45,
                  'Ghana': 14,
                  'Canada': 20,
                  'USA': 88
                 }


Added: converted weight and height into integer datatypes e.g., 10kg -> 10

```python
top_5_country = fifa_dataset['Nationality'].value_counts().head(5)

top_5_country = top_5_country.to_dict()

print(top_5_country)

fifa_dataset = remove_duplicate(remove_newline(fifa_dataset))

#convert weight and height into integer datatypes e.g., 10kg -> 10

def convert_weight(value):
    if isinstance(value, int) or isinstance(value, float):
        return int(value)  # already a number, just return it
    if 'kg' in value:
        return int(value.replace('kg', ''))
    elif 'lbs' in value:
        return round(int(value.replace('lbs', '')) * 0.453592)
    
fifa_dataset['Weight'] = fifa_dataset['Weight'].apply(convert_weight)

if fifa_dataset['Height'].dtype not in ['int64', 'float64']:  # only run if it's not already a number
    fifa_dataset['Height'] = fifa_dataset['Height'].str.extract(r'(\d+)').astype(int)


```

Output:
```
{'England': 1574, 'Germany': 1178, 'Spain': 1019, 'France': 971, 'Argentina': 861}
           ID               Name                      LongName  \
0      158023           L. Messi                  Lionel Messi   
1       20801  Cristiano Ronaldo  C. Ronaldo dos Santos Aveiro   
2      200389           J. Oblak                     Jan Oblak   
3      192985       K. De Bruyne               Kevin De Bruyne   
4      190871          Neymar Jr    Neymar da Silva Santos Jr.   
...       ...                ...                           ...   
18974  247223             Xia Ao                        Ao Xia   
18975  258760           B. Hough                     Ben Hough   
18976  252757        R. McKinley                Ronan McKinley   
18977  243790       Wang Zhen'ao                  Zhen'ao Wang   
18978  252520          Zhou Xiao                     Xiao Zhou   

                                               photoUrl  \
0      https://cdn.sofifa.com/players/158/023/21_60.png   
1     
```

Break the "Long Name" into two columns with first name and last name

```python
fifa_dataset['First Name'] = fifa_dataset['LongName'].str.split(' ').str[0]
fifa_dataset['Last Name'] = fifa_dataset['LongName'].str.split(' ').str[1:].str.join(' ')

# Get current columns as a list
cols = fifa_dataset.columns.tolist()

# Remove First Name and Last Name from wherever they are
cols.remove('First Name')
cols.remove('Last Name')

# Find where LongName is and insert after it
long_name_idx = cols.index('LongName')
cols.insert(long_name_idx + 1, 'First Name')
cols.insert(long_name_idx + 2, 'Last Name')

# Reorder the DataFrame
fifa_dataset = fifa_dataset[cols]

fifa_dataset.head()
```

Output:
```
ID               Name                      LongName First Name  \
0  158023           L. Messi                  Lionel Messi     Lionel   
1   20801  Cristiano Ronaldo  C. Ronaldo dos Santos Aveiro         C.   
2  200389           J. Oblak                     Jan Oblak        Jan   
3  192985       K. De Bruyne               Kevin De Bruyne      Kevin   
4  190871          Neymar Jr    Neymar da Silva Santos Jr.     Neymar   

                   Last Name  \
0                      Messi   
1  Ronaldo dos Santos Aveiro   
2                      Oblak   
3                  De Bruyne   
4        da Silva Santos Jr.   

                                           photoUrl  \
0  https://cdn.sofifa.com/players/158/023/21_60.png   
1  https://cdn.sofifa.com/players/020/801/21_60.png   
2  https://cdn.sofifa.com/players/200/389/21_60.png   
3  https://cdn.sofifa.com/players/192/985/21_60.png   
4  https://cdn.sofifa.com/players/190/871/21_60.png   

                                           p
```

We want to identify outliers in the "Age" column. In this case, we'll use the Standard Deviation (SD) Formula for 4 sds away from the mean

```python
#check for outliers in the age that are 4 standard deviations away from the mean

min_limit = fifa_dataset['Age'].mean() - 4 * fifa_dataset['Age'].std()
max_limit = fifa_dataset['Age'].mean() + 4 * fifa_dataset['Age'].std()

min_limit,max_limit

fifa_dataset[(fifa_dataset['Age'] < max_limit) & (fifa_dataset['Age'] > min_limit)]
```

Output:
```
ID               Name                      LongName First Name  \
0      158023           L. Messi                  Lionel Messi     Lionel   
1       20801  Cristiano Ronaldo  C. Ronaldo dos Santos Aveiro         C.   
2      200389           J. Oblak                     Jan Oblak        Jan   
3      192985       K. De Bruyne               Kevin De Bruyne      Kevin   
4      190871          Neymar Jr    Neymar da Silva Santos Jr.     Neymar   
...       ...                ...                           ...        ...   
18974  247223             Xia Ao                        Ao Xia         Ao   
18975  258760           B. Hough                     Ben Hough        Ben   
18976  252757        R. McKinley                Ronan McKinley      Ronan   
18977  243790       Wang Zhen'ao                  Zhen'ao Wang    Zhen'ao   
18978  252520          Zhou Xiao                     Xiao Zhou       Xiao   

                       Last Name  \
0                          Messi   
1      Ronaldo
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

<!-- BEGIN QUESTION -->

**Question 5:** What insight can you derive from this data?

We've 17920 uique players from 18979
Our players are aged 25 on average with 16 years old being the younges and 53 being the oldest 
England seems to have the highest number of players, followed immediately by Germany and Spain

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True, pdf=False)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:


/Users/chukwuemekaokonkwo/Applications/IDS/ids-fifa-21-player-ratings/venv/lib/python3.14/site-packages/otter/check/utils.py:115: RuntimeWarning: [IPKernelApp] WARNING | Kernel is running over TCP without encryption. All communication (including code and outputs) is sent in plain text and is susceptible to eavesdropping. Use IPC transport or launch with kernel manager-provisioned CurveZMQ keys to enable transport encryption.

  warnings.warn(results.stderr.decode("utf-8"), RuntimeWarning)
<IPython.core.display.HTML object>
```

Tests for q1 (points: 1)
```
>>> fifa_dataset = pd.read_csv('fifa_21_raw_data.csv', engine='python')
>>> 
>>> check_data(fifa_dataset)
False
```


Tests for q2 (points: 1)
```
>>> remove_duplicate(fifa_dataset)
>>> fifa_dataset['Name'].duplicated().any() == False
True
```


Tests for q3 (points: 1)
```
>>> len(fifa_dataset["Name"].unique()) > 1000
True
```
```
>>> type(fifa_dataset["Club"]) == pd.Series
True
```


Tests for q4 (points: 1)
```
>>> type(top_5_country) is dict
True
```
```
>>> len(top_5_country) == 5
True
```
